# STUDENT: `Francisco Villanueva`
## COLLABORATORS: `_YOUR_, _COLLABORATORS_, _HERE_`
* resources:
    * `_YOUR_`
    * `_RESOURCES_`
    * `_HERE_`

# CSCI 3155 Problem Set 3 : Grammars

1. This assignment asks you to write scala programs. There are restrictions associated with how you can solve these problems. Please pay careful heed to those. If you are unsure ask your TA or post on piazza. 
1. Note: `???` indicates that there is a missing function or code fragment that needs to be filled in. In scala, it is also a macro that throws a `NotImplemented` exception. Make sure that you remove the `???` and replace it with the answer. 
1. Use the test cases provided to test your solutions. You are encouraged to write your own tests and you are typically given a space to write these test. **Unless otherwise specified, your personal tests will NOT be graded.** If new cells are made during your development, please delete any extra cells you may have created lest they break our autograder.
1. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).
1. Make sure you fill in any place that says `???`, "`_YOUR_ANSWER _HERE_`", or "`_YOUR_NAME_HERE_`".
1. **Very Important:** Please run the cell that defines the functions `passed`, `testWithMessage`, `testErrorWithMessage` below whenever you restart the notebook.
1. **Very Important:** You are responsible to ensure that your work is submitted correctly. We recommend, after uploading your solution to Canvas, then download the submission and confirm that the correct file was provided for grading.

```scala
// TEST HELPER

// FIRST RUN THIS CELL EVERY TIME YOU START THE NOTEBOOK

// IMPORTS
import scala.util
import scala.annotation.tailrec

// TEST HELPERS
def passed(points: Int) = {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}

/**
 * testWithMessage
 * 
 * test the value returned by calling $e1 against the $expected value.
 * prettify printing of the test results
 *
 * e.g. testWithMessage( () => 1 + 2, 
 *                       3, 
 *                       "my addition test" )
 * 
 * e.g. testWithMessage( () => 1/0, 
 *                       42, 
 *                       "my intentially bad test" )
 *
 * @param e1: a thunk from nothing to T. this is the expression of the thing you really want to test
 * @param expected: this is the expected return value from e1
 * @param testID: the name of the test for pretty printing
 */
def testWithMessage[T](e1: () => T, expected: T, testID: String) = { 
    println(s"Test $testID"); 
    println(s"\t Expected: $expected")
    try {
        val v1: T = e1()
        println(s"\tyour code returned: $v1")
        assert (v1 == expected, s"XXX\tTest $testID FAILED.")
        println(s"\t$testID Passed!")
    } catch {
        case err: Throwable => assert(false, s"XXX\tTest $testID FAILED with unexpected ERROR::\n$err")
    }
}


/**
 * testErrorWithMessage
 * 
 * test the expected error type $expectedErrorTemplate found when executing $e1
 * prettify printing of the test results
 *
 * e.g. val expectedErr = new ArithmeticException
 *      testErrorWithMessage( () => 1/0, 
 *                            expectedErr, 
 *                            "My expected error test")
 *
 * @param e1: a thunk from nothing to T. this is the expression of the thing you really want to test
 * @param expected: this is the expected return value from e1
 * @param testID: the name of the test for pretty printing
 */
def testErrorWithMessage[T, E](e1: () => T, expectedErrorTemplate: E, testID: String) = { 
    println(s"Test $testID"); 
    val expectedErrorClass = expectedErrorTemplate.getClass
    println(s"\tExpected: $expectedErrorClass")
    try {
        val foundValue: T = e1()
        assert(false, s"\tTest $testID FAILED. Expected an error, but found a value")
    } catch {
        case foundError: Throwable => {
            val foundErrorClass = foundError.getClass
            if (foundErrorClass == expectedErrorClass) {
                println(s"\t$testID Passed!")
            } else {
                assert(false, s"XXX\tTest $testID FAILED. Expected error $expectedErrorClass, but found $foundErrorClass.")
            }
        }
    }
}
```

## Problem 1: Grammar to code (35 points)
* Allow us to review the names of the PL staff a bit. Unfortunately there are too many students in the class to do a full student version as we've done in the past. This is a very random grammar that does not hold any significant meaning. 
* convert the following grammar to an object class hierarchy

$$\begin{array}{lll}
Kai & \Rightarrow & Tome \\
&|& Prash \\
&|& Kevin(Ryan,\; Varun) \\
&|& Harsha(Varun) \\
\\
Rahul & \Rightarrow & Tome \\
&|& Shuchi \\
&|& Ryan() \\
&|& Kirby \\
\\
Shuchi & \Rightarrow & Denzil(Tome,\; Rahul) \\
&|& Varun() \\
\\
Tome & \Rightarrow & Spencer(Kai) \\
&|& Nour(Kai,\; Denzil,\; Kai)\\
\end{array}$$

<!-- Spencer & \Rightarrow & Abhishek \\
& | & Alex(Spencer) \\
& | & Ali \\
& | & Andrew \\
& | & Emanuel(Spencer, Abhishek, Varun) \\
& | & Gunn(Spencer, Ali) \\
\\
Abhishek & \Rightarrow & Varun \\
& | & Izzy(Ali, Varun, Ali) \\
& | & Neena(Ali, Ali, Vaurn) \\
& | & Simon \\
\\
Varun & \Rightarrow & Tyler(Spencer) \\
& | & Yohanes \\
\\
Ali & \Rightarrow & Aaron \\
& | & Harley \\
& | & Jack \\
& | & Rahul \\
& | & Kailen \\
& | & Saud \\ -->
* ***Restrictions:*** 
    * You may use `sealed trait`, `case class`, and `case object`
    * If you use standard `class` and `object`, then you must define an `apply` method in order to pass the provided tests for class objects
    * You may NOT use `sealed abstract class` (for grading reasons, not because its bad or unworthy of your studies)

<!-- grade_id: cell-1424445d9b3353f5, points: 0, grade: false, solution: true -->
```scala
// _YOUR_CODE_HERE_
???

// _YOUR_TESTS_HERE
```

<!-- grade_id: cell-721ecf4b51f12766, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
// Kai simple
val x0 = Prash
for (x <- List(x0)) {
    assert(x.isInstanceOf[Kai], "test about Kai")
    assert(!x.isInstanceOf[Rahul], "test about Rahul")
    assert(!x.isInstanceOf[Shuchi], "test about Shuchi")
    assert(!x.isInstanceOf[Tome], "test about Tome")
}
passed(5)
// END TESTS
```

<!-- grade_id: cell-00d8243fcc2c34a1, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
// Rahul simple
val x0 = Kirby
val x1 = Ryan()
for (x <- List(x0, x1)) {
    assert(!x.isInstanceOf[Kai], "test about Kai")
    assert(x.isInstanceOf[Rahul], "test about Rahul")
    assert(!x.isInstanceOf[Shuchi], "test about Shuchi")
    assert(!x.isInstanceOf[Tome], "test about Tome")
}
passed(5)
// END TESTS
```

<!-- grade_id: cell-658085a6d1fc8a01, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
// Shuchi simple
val x0 = Varun()
for (x <- List(x0)) {
    assert(!x.isInstanceOf[Kai], "test about Kai")
    assert(x.isInstanceOf[Rahul], "test about Rahul")
    assert(x.isInstanceOf[Shuchi], "test about Shuchi")
    assert(!x.isInstanceOf[Tome], "test about Tome")
}
passed(5)
// END TESTS
```

<!-- grade_id: cell-3a8f552dc9648092, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
// Tome simplish
val x0 = Spencer(Prash)

for (x <- List(x0)) {
    assert(x.isInstanceOf[Kai], "test about Kai")
    assert(x.isInstanceOf[Rahul], "test about Rahul")
    assert(!x.isInstanceOf[Shuchi], "test about Shuchi")
    assert(x.isInstanceOf[Tome], "test about Tome")
}
passed(5)
// END TESTS
```

<!-- grade_id: cell-dff658b9c722c4c7, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
// Kai simplish
val x0 = Kevin(Ryan(), Varun())
val x1 = Harsha(Varun())
for (x <- List(x0, x1)) {
    assert(x.isInstanceOf[Kai], "test about Kai")
    assert(!x.isInstanceOf[Rahul], "test about Rahul")
    assert(!x.isInstanceOf[Shuchi], "test about Shuchi")
    assert(!x.isInstanceOf[Tome], "test about Tome")
}
passed(5)
// END TESTS
```

<!-- grade_id: cell-feff45d5966e3dd1, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
// integrations Shuchi
val x0 = Denzil(Spencer(Prash), Kirby)
val x1 = Denzil(Nour(Prash, x0, Harsha(Varun())), Ryan())

for (x <- List(x0, x1)) {
    assert(!x.isInstanceOf[Kai], "test about Kai")
    assert(x.isInstanceOf[Rahul], "test about Rahul")
    assert(x.isInstanceOf[Shuchi], "test about Shuchi")
    assert(!x.isInstanceOf[Tome], "test about Tome")
}
passed(5)
// END TESTS
```

<!-- grade_id: cell-6380ac4e50b6de04, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
// integrations tome
val x0 = Spencer(Kevin(Ryan(), Varun()))
val x1 = Nour(Prash, Denzil(x0, Ryan()), Harsha(Varun()))

for (x <- List(x0, x1)) {
    assert(x.isInstanceOf[Kai], "test about Kai")
    assert(x.isInstanceOf[Rahul], "test about Rahul")
    assert(!x.isInstanceOf[Shuchi], "test about Shuchi")
    assert(x.isInstanceOf[Tome], "test about Tome")
}
passed(5)
// END TESTS
```

# Problem 2 Custom List of Custom Tuple (65 points)
Throughout the course we'll be working with pattern matching and unpacking interesting data structures. Let us learn more about lists and tuples now to hopefully shed some light on this matter.

## Problem 2.A. unpacking Scala List and Tuple (10 points)
We've learned that `5 +: List(20, 7)` constructs a list `List(5, 20, 7)`, and that pattern `_ +: _` can be used to unpack the list as well? Run the following code and note the value stored at `x` and `y`.

```scala
val x +: y = List(5, 20, 7)
```

We know that `(22, 14)` constructs a tuple, and that pattern `(_, _)` can be used to unpack a tuple as well. Run the following code and note the value stored at `x` and `y`.

```scala
val (x, y)  = (22, 14) 
```

We know that `List(1,2,3)` constructs a list, can we use `val List(_,_,_)` to unpack the list as well? Run the following code and note the value stored at `x`,`y`, and `z`.

```scala
val List(x,y,z)  = List(5,10,15)
```

If we aren't certain about the length of the list, we might find the `val _ +: _` pattern easier. Run the following code and note the value stored at `first`, and `rest`.

```scala
val first +: rest = List(5,10,15)
```

But then what about a more complex list where we care about some element of the list, but also values within that element? The `_@_` pattern allows us to name the thing that we will further match on. Run the following code and note the value stored in `second`, `foo`, and `bar`.

```scala
val _ +: ( second@List(foo, bar) ) +: _ = List( List(1,2), List(30, 40), List(987,654))
```

Let us use the above concepts to unpack some lists to construct a few special variables. Your code should unpack the `listOfNumbers`, `aTuple`, and `listOfTupleStringNumber` and construct the following values in the following variables:
* from listOfNumbers
    * alpha = 5
    * bravo = 22
* from aTuple
    * charlie = "hi"
    * delta = "bye"
* from listOfTupleStringNumber
    * echo = ("x", 10)
    * foxtrot = "x"
    * gamma = List(("y", 20))
    * hotel = 20

You can use multiple lines and extra variables to achieve this, but we encourage that you challenge yourself to use `@` symbol and/or `_` in your solution to ignore values that aren't important to you. <a href="https://www.baeldung.com/scala/pattern-matching">This Baeldung tutorial</a> demonstrates how to do this in a pattern match statement. While we won't use `match` just yet, the concept for unpacking remains the same.

```scala
// PROVIDED (remember to run this cell)
val listOfNumbers = List(5, 42, 22)
val aTuple = ("hi", "bye")
val listOfTupleStringNumber = List(("x", 10), ("y", 20))
```

<!-- grade_id: cell-d8e8896d71681223, points: 0, grade: false, solution: true -->
```scala
// _YOUR_CODE_HERE_
???

// _YOUR_TESTS_HERE_
```

<!-- grade_id: cell-85cbcf63a605294a, points: 1, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => alpha, 5, "ALPHA")
passed(1)
// END TESTS
```

<!-- grade_id: cell-0fdf7f6087dae4c5, points: 1, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => bravo, 22, "BRAVO")
passed(1)
// END TESTS
```

<!-- grade_id: cell-3b6c5272488589a5, points: 1, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => charlie, "hi", "CHARLIE")
passed(1)
// END TESTS
```

<!-- grade_id: cell-eb9556b0a5d9d5b7, points: 1, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => delta, "bye", "DELTA")
passed(1)
// END TESTS
```

<!-- grade_id: cell-60edc98be5b4f90e, points: 2, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => echo, ("x", 10), "ECHO")
passed(2)
// END TESTS
```

<!-- grade_id: cell-50ea18039a3a876d, points: 1, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => foxtrot, "x", "FOXTROT")
passed(1)
// END TESTS
```

<!-- grade_id: cell-a40b91a4d6e4e7bd, points: 2, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => gamma, List(("y", 20)), "GAMMA")
passed(2)
// END TESTS
```

<!-- grade_id: cell-ef71b70bb7bae53d, points: 1, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => hotel, 20, "HOTEL")
passed(1)
// END TESTS
```

## Problem 2.B. constructing custom list of custom tuple (15 points)
Consider the following grammar that defines a custom list over custom tuples.

$$\begin{array}
\ MyTupleList & \Rightarrow & EmptyList \\
&|& MyListNode(Element,\; MyTupleList) \\
\\
Element & \Rightarrow & MyTupleNode(d, d) \\
\\ 
&& \texttt{d is a Scala Double}
\end{array}$$

Please write this as Scala code?

<!-- grade_id: cell-11ffdd200e182e63, points: 0, grade: false, solution: true -->
```scala
// _YOUR_CODE_HERE_
???


// _YOUR_TESTS_HERE_
```

<!-- grade_id: cell-433266190293384d, points: 10, grade: true, solution: false -->
```scala
// BEGIN TESTS
val x0 = EmptyList
val x1 = MyListNode(MyTupleNode(1,2), x0)
for (x <- List(x0, x1)) {
    assert(x.isInstanceOf[MyTupleList])
    assert(!x.isInstanceOf[Element])
    assert(!x.isInstanceOf[Double])
}
passed(10)
// END TESTS
```

<!-- grade_id: cell-e4174d2cd27c8cb2, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
val x0 = MyTupleNode(1, 2)
val x1 = MyTupleNode(15, 10)
for (x <- List(x0, x1)) {
    assert(!x.isInstanceOf[MyTupleList])
    assert(x.isInstanceOf[Element])
    assert(!x.isInstanceOf[Double])
}
passed(5)
// END TESTS
```

## Problem 2.C. unpacking custom list of custom tuple (20 point)

Unpack variable `foo` representing the Scala List `List((5, 10), (25, 7), (1000, 20000))` using our custom data structure, unpack `foo` such that the following holds true:
* india = 10
* juliette = 25
* kilo = MyTupleNode(1000, 20000)
* mike = 1000


You can use multiple lines and extra variables to achieve this, but we encourage that you challenge yourself to use `@` symbol and/or `_` in your solution to ignore values that aren't important to you. <a href="https://www.baeldung.com/scala/pattern-matching">This Baeldung tutorial</a> demonstrates how to do this in a pattern match statement. While we won't use `match` just yet, the concept for unpacking remains the same.

```scala
// GIVEN
// List((5, 10), (25, 7), (1000, 20000))
val foo = MyListNode(MyTupleNode(5, 10), 
                   MyListNode(MyTupleNode(25, 7),
                              MyListNode(MyTupleNode(1000, 20000),
                                         EmptyList)))
```

<!-- grade_id: cell-6419948b8fab433e, points: 0, grade: false, solution: true -->
```scala
// _YOUR_CODE_HERE_
???
```

<!-- grade_id: cell-d221fe20a87b7596, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => india, 10, "INDIA")
passed(5)
// END TESTS
```

<!-- grade_id: cell-afef679e869a44e3, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => juliette, 25, "JULIETTE")
passed(5)
// END TESTS
```

<!-- grade_id: cell-f8b26c57f30ad370, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => kilo, MyTupleNode(1000, 20000), "KILO")
passed(5)
// END TESTS
```

<!-- grade_id: cell-fecd78598835ce5b, points: 5, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => mike, 1000, "MIKE")
passed(5)
// END TESTS
```

## Problem 2.D. Pattern Matching on custom list of custom tuple (20 points)
Below we demonstrate how to solve 'sum' recursively using pattern matching on a Scala `List[(Double, Double)]`. Your task is to solve the same function using pattern matching but on a `MyTupleList` instead. A few code stubs are provided for you to choose from if you'd like.

* ***Restrictions:***
    * Solution must be recursive (optionally tail recursive including the `@tailrec` decorator)
    * Function must be named `sumMyTupleList`
    * Function must have a default function type `MyTupleList => Double`
    * Function cannot call `sum`
    * Solution must use a pattern `match` statement
    * Can NOT use the keywords `return`, `var`, `for`, or `while`

```scala
// GIVEN
def sum(lst: List[(Double, Double)]): Double = {
    lst match {
        case List() => 0
        case (firstLeft, firstRight) +: rest => firstLeft + firstRight + sum(rest)
    }
}
```

<!-- grade_id: cell-e4c22116e854e7ca, points: 0, grade: false, solution: true -->
```scala
// choose any of these stubs, or make your own that passes the tests

// def sumMyTupleList(lst: MyTupleList): Double = { ...

// @tailrec
// def sumMyTupleList(l: MyTupleList, acc: Double = 0): Double = { ...

// def sumMyTupleList(l: MyTupleList): Double = { ...<some kind of helper that is tail recursive>...

// _YOUR_CODE_HERE_
???


// _YOUR_TESTS_HERE_
```

<!-- grade_id: cell-4f435cf7753a9bcd, points: 3, grade: true, solution: false -->
```scala
// BEGIN TESTS
testWithMessage(() => sumMyTupleList(EmptyList), 0, "test sumMyTupleList base case")
passed(3)
// END TESTS
```

<!-- grade_id: cell-4f7ad81cf09baebb, points: 7, grade: true, solution: false -->
```scala
// BEGIN TESTS
val t1 = MyTupleNode(5, 10)
val l1 = MyListNode(t1, EmptyList)
testWithMessage(() => sumMyTupleList(l1), 15, "test sumMyTupleList base case")

val rand = new scala.util.Random
val d1 = rand.nextDouble()
val d2 = rand.nextDouble()
val t2 = MyTupleNode(d1, d2)
val l2 = MyListNode(t2, EmptyList)
testWithMessage(() => sumMyTupleList(l2), d1 + d2, "test sumMyTupleList base case")
passed(7)
// END TESTS
```

<!-- grade_id: cell-91c28f946b83fb17, points: 10, grade: true, solution: false -->
```scala
// BEGIN TESTS
val foo = MyListNode(MyTupleNode(5, 10), 
                   MyListNode(MyTupleNode(25, 7),
                              MyListNode(MyTupleNode(1000, 20000),
                                         EmptyList)))

val bar = MyListNode(MyTupleNode(7, 12), 
                   MyListNode(MyTupleNode(27, 9),
                              MyListNode(MyTupleNode(1002, 20002),
                                         EmptyList)))

val baz = MyListNode(MyTupleNode(10, 20), 
                   MyListNode(MyTupleNode(50, 14),
                              MyListNode(MyTupleNode(2000, 40000),
                                         EmptyList)))

testWithMessage(() => sumMyTupleList(foo), 21047, "test sumMyTupleList 1")
testWithMessage(() => sumMyTupleList(bar), 21047+12, "test sumMyTupleList 2")
testWithMessage(() => sumMyTupleList(baz), 21047*2, "test sumMyTupleList 3")
passed(10)
// END TESTS
```

## Challenge Problem (0 points)
* **Challenge problems are optional and not for credit.**
* For Scala `List[(Int, Int)]` and/or `MyTupleList` implement each of the following and test your work
* a function that finds the value of the product of all numbers in the collection (lookup Scala `foldRight`, `foldLeft`, `reduce`). 
* a function that finds the largest number in the collection (lookup Scala `foldRight`, `foldLeft`, `reduce`).
* a function that confirms if some specific number exists in the collection (lookup Scala `contains`).
* a function that confirms if some number exists in the collection such that some condition holds true e.g. has an even values (lookup Scala `filter`).
* a function that modifies each element of the collection in some way e.g. increment each number (lookup Scala `map`).

<!-- grade_id: cell-b1aaed0815ce24d2, points: 0, grade: false, solution: true -->
```scala
// _YOUR_CODE_HERE_
???
```

## That's all folks